# CS211: Data Privacy
## Homework 4

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

def laplace_mech(v, sensitivity, epsilon):
    return v + np.random.laplace(loc=0, scale=sensitivity / epsilon)

def pct_error(orig, priv):
    return np.abs(orig - priv)/orig * 100.0

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

```python
adult['fnlwgt']
```

Output:
```
0         77516
1         83311
2        215646
3        234721
4        338409
          ...  
32556    257302
32557    154374
32558    151910
32559    201490
32560    287927
Name: fnlwgt, Length: 32561, dtype: int64
```

## Question 1 (10 points)

Complete the definition of `dp_sum_capgain` below. Your definition should compute a differentially private sum of the "Capital Gain" column of the `adult` dataset, and have a total privacy cost of `epsilon`.

<!-- grade_id: cell-6a369f61d8eef943, points: 0, grade: false, solution: true -->
```python
def dp_sum_capgain(epsilon):
    sensitivity = 100000
    return laplace_mech(adult['Capital Gain'].clip(upper=sensitivity).sum(), sensitivity, epsilon)
dp_sum_capgain(1.0)
```

Output:
```
35316454.5521916
```

<!-- grade_id: cell-e5b1229ba2249f71, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 1

real_sum = adult['Capital Gain'].sum()
r1 = np.mean([pct_error(real_sum, dp_sum_capgain(0.1)) for _ in range(100)])
r2 = np.mean([pct_error(real_sum, dp_sum_capgain(1.0)) for _ in range(100)])
r3 = np.mean([pct_error(real_sum, dp_sum_capgain(10.0)) for _ in range(100)])

print("Average errors:", r1, r2, r3)

assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 10
assert r2 < 2
assert r3 < 0.2
```

Output:
```
Average errors: 9.178603376279346 9.035920857170886 9.062429982454166
AssertionError                            Traceback (most recent call last)
/tmp/ipykernel_146424/3302739200.py in <module>
     12 assert r3 > 0
     13 assert r1 < 10
---> 14 assert r2 < 2
     15 assert r3 < 0.2

AssertionError:
```

## Question 2 (10 points)

In 2-5 sentences each, answer the following:

- What clipping parameter did you use in your definition of `dp_sum_capital`, and why?
- What was the sensitivity of the query you used in `dp_sum_capital`, and how is it bounded?
- Argue that your definition of `dp_sum_capital` has a total privacy cost of `epsilon`

<!-- grade_id: cell-986eec754f32c072, points: 10, grade: true, solution: true -->
100000, because that was a reasonable upper bound on the sizes in the dataset.
The sensitivity is 100000 as well, because the sensitivity is bounded by the largest possible size in the dataset, so if we clip anything above that, we ensure that the function cannot increase or decrease by >100000 with the addition of more data.
Only one query is made with bounded sensitivity using the laplace mechanism, the sum of the one cost epsilon is epsilon.

## Question 3 (10 points)

Complete the definition of `dp_avg_capgain` below. Your definition should compute a differentially private average (mean) of the "Capital Gain" column of the adult dataset, and have a **total privacy cost of epsilon**.

<!-- grade_id: cell-93f2a5153dca0e95, points: 0, grade: false, solution: true -->
```python
def dp_avg_capgain(epsilon):
    return dp_sum_capgain(epsilon / 2) / laplace_mech(adult['Capital Gain'].shape[0], 1, epsilon / 2)

dp_avg_capgain(1.0)
```

Output:
```
1086.0638347435206
```

<!-- grade_id: cell-ef77b325b5c58908, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 3

real_avg = adult['Capital Gain'].mean()
r1 = np.mean([pct_error(real_avg, dp_avg_capgain(0.1)) for _ in range(100)])
r2 = np.mean([pct_error(real_avg, dp_avg_capgain(1.0)) for _ in range(100)])
r3 = np.mean([pct_error(real_avg, dp_avg_capgain(10.0)) for _ in range(100)])

print("Average errors:", r1, r2, r3)

assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 20
assert r2 < 4
assert r3 < 0.4
```

Output:
```
Average errors: 5.779807314170563 0.6453845009574745 0.05328625228506038
```

## Question 4 (10 points)

In 2-5 sentences each, answer the following:

- Argue that your definition of `dp_avg_capgain` has a total privacy cost of `epsilon`
- For sums and averages, which seems to be more important for accuracy - the value of the clipping parameter $b$ or the scale of the noise added? Why?
- Do you think the answer to the previous point will be true for every dataset? Why or why not?

<!-- grade_id: cell-7ea16016366529bd, points: 10, grade: true, solution: true -->
I divided it into a sum and a count query each with epsilon = `epsilon/2`. By sequential composition, total privacy cost is `epsilon`.
From playing around, it looks like variation in the clipping parameter is more problematic because this dataset is skewed such that the majority of people have very little capital gains, but those who have more than a little tend to have quite a bit. So when we clip on either end, we wind up much more inaccurate than if we had only a few outliers. Therefore, I don't think this property necessarily generalizes.

## Question 5 (20 points)

Write a function `auto_avg` that returns the differentially private average of a Pandas series `s`. Your function should automatically determine the clipping parameter `b`, and should enforce differential privacy for a **total privacy cost** of `epsilon`. You can assume that all values are non-negative (i.e. 0 or greater).

<!-- grade_id: cell-10b352b38a9b5f85, points: 0, grade: false, solution: true -->
```python
def auto_avg(series, epsilon):
    sub_e = epsilon / 16
    last = 0
    sensitivity = 0
    for i in range(15):
        sensitivity = 3 ** i
        new = laplace_mech(series.clip(upper=sensitivity).sum(), sensitivity, sub_e)
        if abs(new - last) < (new * 0.15):
            print(f"b = {sensitivity}")
            return new / laplace_mech(series.shape[0], 1, sub_e)
        last = new
    return last / laplace_mech(series.shape[0], 1, sub_e)
```

<!-- grade_id: cell-49fb57b1ebfcc0da, points: 20, grade: true, solution: false -->
```python
# TEST CASE for question 5
r1 = np.mean([pct_error(adult['Age'].mean(), auto_avg(adult['Age'], 1.0)) for _ in range(20)])
print('capital')
r2 = np.mean([pct_error(adult['Capital Gain'].mean(), auto_avg(adult['Capital Gain'], 1.0)) for _ in range(20)])
print('fnlwgt')
r3 = np.mean([pct_error(adult['fnlwgt'].mean(), auto_avg(adult['fnlwgt'], 1.0)) for _ in range(20)])

print('Average errors:', r1, r2, r3)
assert r1 > 0
assert r2 > 0
assert r3 > 0
assert r1 < 1
assert r2 < 100
assert r3 < 1
```

Output:
```
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
b = 243
capital
b = 177147
b = 531441
b = 531441
b = 4782969
b = 177147
b = 177147
b = 177147
b = 1594323
b = 531441
b = 531441
b = 531441
b = 1594323
b = 531441
b = 177147
b = 177147
b = 1594323
b = 531441
fnlwgt
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
b = 1594323
Average errors: 0.21577032731995588 21.97406407533915 0.29386529629599006
```

## Question 6

In 2-5 sentences each, answer the following:

- Explain your strategy for implementing `auto_avg`
- Argue informally that your definition of `auto_avg` has a total privacy cost of `epsilon`
- Did your solution work well for all three example columns? If it did not work well on any of them, why not?
- When is your solution likely to *not* work well? (i.e. what properties does the data have to have, in order for your solution to not work well?)

<!-- grade_id: cell-0ff63739ed66adf7, points: 10, grade: true, solution: true -->
Tried to implement logarithmic strategy from textbook, tweaking parameters landed me at powers of 3, given it needs to cover a wide array of data sizes.
By sequential composition, my method runs at most 16 queries with an epsilon of `epsilon/16`. (15 checks to find b, and one count query.) which sum to at most `epsilon` and at least `epsilon/8`
It worked for all of them, but I did spend some time tweaking the parameters to get it there, so not necessarily generalizable.
When data exceeds 3 ** 14, or requires clipping with more precision, this solution will break down.