# CS211: Data Privacy
## Homework 1

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
adult_pii = adult[['Name', 'DOB', 'SSN', 'Zip', 'Age']]
adult_deid = adult.drop(columns=['Name', 'SSN'])
```

## Question 1 (20 points)

Using the dataframes `adult_pii` and `adult_deid`, write code to conduct a linking attack to recover the names of as many individuals in `adult_deid` as possible. Your solution should be parameterized by the set of columns to use in the attack.

<!-- grade_id: cell-e6389ceace711c32, points: 0, grade: false, solution: true -->
```python
def linking_attack(cols):
    return adult_deid.merge(adult, on=cols)

```

<!-- grade_id: cell-251b9f1c6f7b4630, points: 20, grade: true, solution: false -->
```python
# TEST CASES for Question 1

assert len(linking_attack(['Zip'])) == 43191
assert len(linking_attack(['Zip', 'DOB'])) == 32563
assert len(linking_attack(['Zip', 'Age'])) == 32755
```

## Question 2 (5 points)

How many individuals in this dataset are uniquely identified by their Zip code? How many are uniquely identified by their age?

Hint: note that the number of *unique ZIP codes* is **different** from the number of *individuals uniquely identified by ZIP code*.

Hint: you can use the `value_counts` method to count the number of occurences of each value in a series.

```python
def unique_zipcode():
    return len(adult.value_counts("Zip").loc[lambda x: x == 1])

def unique_dob():
    # YOUR CODE HERE
    return len(adult.value_counts("DOB").loc[lambda x: x == 1])
```

```python
# TEST CASES (INCOMPLETE) for Question 2

assert unique_zipcode() > 0
assert unique_dob() > 0

```

Output:
```
Zip
68137    1
73284    1
68141    1
73291    1
71324    1
        ..
36771    1
36765    1
36761    1
36758    1
99999    1
Length: 23513, dtype: int64
```

## Question 3 (10 points)

Write code to determine the `Education-Num` of the individual named Ardyce Golby by performing a differencing attack. Your code should *only* use aggregate data to find Ardyce's education number.

<!-- grade_id: cell-ef28e136429ae8c2, points: 0, grade: false, solution: true -->
```python
def ardyce_education():
    return adult['Education-Num'].sum() - adult[adult['Name'] != 'Ardyce Golby']['Education-Num'].sum()
```

```python
# TEST CASE for Question 3
assert ardyce_education() == 12
```

```python

```