# CS211: Data Privacy
## In-Class Exercises, week of 9/12/2022

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

Output:
```
/home/lxi/anaconda3/envs/ten/lib/python3.10/site-packages/scipy/__init__.py:146: UserWarning: A NumPy version >=1.16.5 and <1.23.0 is required for this version of SciPy (detected version 1.23.1
  warnings.warn(f"A NumPy version >={np_minversion} and <{np_maxversion}"
```

## Question 1

Write a *counting query* to determine whether or not Karrie Trusslove's age is 39.

```python
F_2 = float(input('Enter temperature in degrees F: '))

F_1 = (F_2 - 10)

F_3 = (F_2 + 10)

C_1 = ((5) / (9) * (F_1 - 32))

C_2 = ((5) / (9) * (F_2 - 32))

C_3 = ((5) / (9) * (F_3 - 32))


print(f'{"Fahrenheit":<5} {"Celsius":>5}')

print(f'{F_1:<5} {C_1:>10.4f}')

print(f'{F_2:<5} {C_2:>10.4f}')

print(f'{F_3:<5} {C_3:>10.4f}')
```

Output:
```
Fahrenheit Celsius
61.0     16.1111
71.0     21.6667
81.0     27.2222
```

<!-- grade_id: cell-e9c0c06eb7735ea3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def karrie_query():
    karries_row = adult[adult['Name'] == 'Karrie Trusslove']
    return karries_row[karries_row['Age'] == 39].shape[0]
```

<!-- grade_id: cell-f93cfd0eb59130ed, points: 5, grade: true, solution: false -->
```python
# TEST CASE 

assert karrie_query() == 1
```

## Question 2

Add Laplace noise to the counting query you wrote in the last question to ensure differential privacy for $\epsilon = 1.0$.

<!-- grade_id: cell-8ad72c55497bc534, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
epsilon = 0.0001
sensitivity = 1
def dp_karrie_query():
    return karrie_query() + np.random.laplace(loc=1, scale=sensitivity/epsilon)
```

<!-- grade_id: cell-a7c9bd084e5c5bdf, points: 5, grade: true, solution: false -->
```python
# TEST CASE 

q2_runs = [dp_karrie_query() for _ in range(100)]
noise_runs = [np.random.laplace(loc=1, scale=1/epsilon) for _ in range(100)]

assert stats.wasserstein_distance(q2_runs, noise_runs) < 1
```

## Question 3

In 2-4 sentences, describe how `dp_karrie_query` protects Karrie's privacy.

<!-- grade_id: cell-07337a085c0ba0e1, points: 5, grade: true, solution: true -->
We add noise to introduce artificial uncertainty to the answer
The noise prevents us from making inferences about Karrie's age with 100% accuracy
I am roughly equally likely to see a particular answer for the query, regardless of carrie's age

## Question 4

Implement a function to construct an *empirical distribution* over random samples from the Laplace distribution with location 0 and scale `scale`. Use 1000 samples.

<!-- grade_id: cell-e7ddf0e4b195dfdc, points: 0, grade: false, solution: true -->
```python
def empirical_dist_lap(scale):
    # YOUR CODE HERE
    raise NotImplementedError()

epsilon1 = 1.0
epsilon2 = 2.0

plt.hist(empirical_dist_lap(1/epsilon1), bins=50);
plt.hist(empirical_dist_lap(1/epsilon2), bins=50, alpha=.7);
```

<!-- grade_id: cell-e68891bb70939fd0, points: 5, grade: true, solution: false -->
```python
# PLACEHOLDER for test cases

```

```python
np.exp(1)
```

```python
counts1, bins, _ = plt.hist([np.random.laplace(loc=0, scale=1/1) for _ in range(100000)], bins=10)
counts2, _, _ = plt.hist([np.random.laplace(loc=1, scale=1/1) for _ in range(100000)], bins=bins, alpha=.5)
list(zip(bins, counts1, counts2))

for bin_val, count1, count2 in zip(bins,counts1,counts2):
    print(bin_val, count1, count2, count1 / count2)
```

## Question 5

Consider the program below.

```python
counting_query = lambda: 1

result_1 = counting_query() + np.random.laplace(loc=0, scale=1/epsilon1)
result_2 = counting_query() + np.random.laplace(loc=0, scale=1/epsilon2)
result_3 = counting_query() + np.random.laplace(loc=0, scale=1/epsilon1)

print(result_1, result_2, result_3)
```

Write a function `total_epsilon` that returns the total privacy cost for this program.

<!-- grade_id: cell-32b1cdb0a2990e8f, points: 0, grade: false, solution: true -->
```python
def total_epsilon():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e484920fc1429ad0, points: 5, grade: true, solution: false -->
```python
# Placeholder for test case

```