# CS211: Data Privacy
## In-Class Exercises, week of 9/5/2022

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs211-data-privacy/raw/master/homework/adult_with_pii.csv')
```

## Question 1

Write code to group the `adult` dataset by a single column and count the number of members in each group.

<!-- grade_id: cell-055399c92e676783, points: 0, grade: false, solution: true -->
```python
def group_one_count(col):
    return adult.value_counts([col])
```

```python
group_one_count('Education')
```

Output:
```
Education   
HS-grad         10501
Some-college     7291
Bachelors        5355
Masters          1723
Assoc-voc        1382
11th             1175
Assoc-acdm       1067
10th              933
7th-8th           646
Prof-school       576
9th               514
12th              433
Doctorate         413
5th-6th           333
1st-4th           168
Preschool          51
dtype: int64
```

<!-- grade_id: cell-befc11c8be0bf91b, points: 5, grade: true, solution: false -->
```python
s = group_one_count('Education')
assert s['10th'] == 933
assert s['9th'] == 514
assert s['Some-college'] == 7291
```

## Question 2

Write code to group the `adult` dataset by two columns and count the number of members in each group.

<!-- grade_id: cell-f6a37efe405c4a98, points: 0, grade: false, solution: true -->
```python
def group_two_count(col1, col2):
    return adult.value_counts([col1, col2])
```

```python
group_two_count('Occupation', 'Education')
```

Output:
```
Occupation         Education   
Craft-repair       HS-grad         1922
Prof-specialty     Bachelors       1495
Exec-managerial    Bachelors       1369
Adm-clerical       HS-grad         1365
                   Some-college    1281
                                   ... 
Machine-op-inspct  Prof-school        1
Protective-serv    1st-4th            1
Prof-specialty     5th-6th            1
                   Preschool          1
Protective-serv    5th-6th            1
Length: 201, dtype: int64
```

<!-- grade_id: cell-88caedc5d6e44568, points: 5, grade: true, solution: false -->
```python
s = group_two_count('Occupation', 'Education')
assert s['Transport-moving']['Doctorate'] == 1
assert s['Adm-clerical']['10th'] == 38
```

## Question 3

Write code to perform a differencing attack to determine Brenn McNeely's age using only the `mean` aggregation function over large groups.

<!-- grade_id: cell-6fb63040052b0857, points: 0, grade: false, solution: true -->
```python
def get_brenns_age():
    mean = adult["Age"].mean()
    mean_no_brenn = adult[adult['Name'] != 'Brenn McNeely']["Age"].mean()
    return (len(adult)) * mean - len(adult[adult['Name'] != 'Brenn McNeely']) * mean_no_brenn

brenns_age = get_brenns_age()
```

<!-- grade_id: cell-d66b7d6c32d256c6, points: 5, grade: true, solution: false -->
```python
assert brenns_age == 38.0
```

```python
df = adult[['Education', 'Marital Status', 'Target']]
df.head(10)
df.value_counts().sort_values()
```

Output:
```
Education     Marital Status         Target
5th-6th       Never-married          >50K         1
Preschool     Divorced               <=50K        1
12th          Widowed                >50K         1
7th-8th       Married-spouse-absent  >50K         1
12th          Never-married          >50K         1
                                               ... 
Bachelors     Never-married          <=50K     1613
              Married-civ-spouse     >50K      1855
Some-college  Never-married          <=50K     2880
HS-grad       Never-married          <=50K     3043
              Married-civ-spouse     <=50K     3336
Length: 173, dtype: int64
```

## Question 4

Does the dataset `df` satisfy $k$-Anonymity for $k=3$? Why not?

<!-- grade_id: cell-a9793d100ba20894, points: 5, grade: true, solution: true -->
No, because there's many combinations, such as 12th/never-married/>50k which belong to just one person.

## Question 5

What columns should we designate as quasi-identifiers in the dataset `df`, and why?

<!-- grade_id: cell-5cfd666564429892, points: 5, grade: true, solution: true -->
In general we should mark every column we want to keep as a QI. We should sacrifice data which is irrelevant to the question we are asking first, to preserve the information which we are trying to glean something from.

## Question 6

Imagine the column `Target` is *not* a quasi-identifier, and we generalize the dataset to achieve $k$-Anonymity for $k=2$ as follows:

- Replace each education level below "HS Grad" with `< HS` and others with `>= HS`
- Replace marital status with `Married` and `Not Married`
- Delete rows as required to achieve $k$-anonymity for $k=2$

For which rows is a homogeneity attack possible, and why?

<!-- grade_id: cell-5932f02f09b5cacb, points: 5, grade: true, solution: true -->
To the best of my understanding, if there exists a combination of education and marital status for which all people fall into the same Target category, then a homogeneity attack is possible.

## Question 7

Write code to generalize the `Zip` column in the adult dataset.

<!-- grade_id: cell-be9a4264a0e58938, points: 0, grade: false, solution: true -->
```python
def generalize_zip(zip, l):
    return zip - (zip % (10 ** l))
```

```python
# Test cases
assert generalize_zip(47401, 0) == 47401
assert generalize_zip(47401, 1) == 47400
assert generalize_zip(47401, 2) == 47400
assert generalize_zip(47401, 3) == 47000
assert generalize_zip(47401, 4) == 40000
assert generalize_zip(47401, 5) == 0
```

```python

```